# Exercise 2 - Lists and indices

The exercise for this week is meant to help you better understand data types and lists in Python, and practice saving changes to your files using Git and GitHub. 

Exercise 2 consists of 3 problems that you can find in this notebook. For each problem you need to modify the notebook either by fixing problems in existing code, or by adding your own solutions. Remember to push your changes to GitHub already while working, not only at the very end!

- **Exercise 2 is due by 09:00 on 18.9.**
- Scores on this exercise are out of **20 points**.


## Where to find help

If you're having trouble, please remember you can

- Come to the work sessions on Thursday and/or Friday for help from the course assistants
- Check out [the hints for this week's exercise](https://geo-python.github.io/site/lessons/L2/exercise-2.html) on the course webpage
- Post questions to Slack

## Problem 0 - Cloning Exercise repository to JupyterLab (*0 points*)

Before starting to work with the actual problems for this week, you should start a new JupyterLab instance and clone **your own** Exercise 2 repository (e.g. `exercise-2-htenkanen`) into the instance using **Git** as we saw in the [lecture this week](https://geo-python.github.io/site/lessons/L2/git-basics.html). 

## Problem 1 - Making changes to code and using Git (*7 points*)

Your first task for this week is to make some modifications to the Python code in the cells below to fix the errors it gives once you run the code cells.
The code should allow users to find the operational time of an [FMI observation station](http://en.ilmatieteenlaitos.fi/observation-stations) by setting the `selected_station` variable.
*Don't worry about the case of a station name that is not on the list.*

**The main aim is to print out an informative sentence that tells how long the selected station has been operational, for example** `'The Helsinki Harmaja station has been operational for 30 years.'`. However, if you try to run all the code cells under Problem 1, you will run into errors. Your task is to find **3 things to fix** in the code cells under Problem 1. Once you succeed, the final print statement should work.

**NOTE: Some of the cells in this Jupyter Notebook are "Read-only" which means that you are not able to edit them.**

### Grading for Problem 1

Your score for this problem will be based on

- Fixing **3 major code problems** to get the code working as expected
- Committing each change separately to GitHub using Git
- Listing the changes you needed to make to get the code working in plain English by modifying the existing text or adding a Markdown cell above each change

Here we create and fill lists of station names and starting years for observation.
**NOTE**: Long lines can be split into multiple lines using the "\\" character. Split lines can be indented freely. 

*These code cells are "Read-only" - you won't be able to modify the input lists.*

```python
# Read-only code cell
station_names = ['Harmaja', 'Kaisaniemi', 'Kaivopuisto', 'Kumpula', 'lighthouse', \
                'Malmi airfield', 'Suomenlinna aaltopoiju', 'Vuosaari harbour']
```

```python
# Read-only code cell
station_start_years = [1989, 1844, 1904, 2005, 2003, 1937, 2016, 2012]
```

Below we can define the selected station. 

*you can start modifying the code starting from the code cell below.*

<!-- grade_id: selected_station, points: 0, grade: false, solution: true -->
```python
selected_station = None
### BEGIN SOLUTION
selected_station = 'Harmaja'
### END SOLUTION
```

<!-- grade_id: selected_station_test, points: 2, grade: true, solution: false -->
```python
print(selected_station)

### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal

# Validate that selected_station is an element in the list station_names
ok_(selected_station in station_names, \
    'selected_station = '+str(selected_station)+' is not an element in station_names')
### END HIDDEN TESTS
```

Output:
```
Harmaja
```

Next we find the location of the selected station in the list of station names.

```python
# Read-only code cell
station_index = station_names.index(selected_station)
```

Now we can calculate the years the station has been operational.

<!-- grade_id: station_years, points: 0, grade: false, solution: true -->
```python
station_years = 2019 - station_start_years[selected_station]
### BEGIN SOLUTION
station_years = 2019 - station_start_years[station_index]
### END SOLUTION
```

Finally, we print the station name and number of years of operation to the screen.
A sample print statement on screen is 'The Helsinki Harmaja station has been operational for 30 years.'

<!-- grade_id: print_statement, points: 0, grade: false, solution: true -->
```python
print_statement = "The Helsinki " + selected_station + " station has been operational for "+ \
      str(station_years)+ " years.

### BEGIN SOLUTION
selected_station = "The Helsinki " + selected_station + " station has been operational for "+ \
      str(station_years)+ " years."
### END SOLUTION

print(print_statement)
```

Output:
```
The Helsinki Harmaja station has been operational for 30 years.
```

<!-- grade_id: print_statement_test, points: 1, grade: true, solution: false -->
```python
#Validate that the print statement is correct
assert_equal(print_statement, 'The Helsinki Harmaja station has been operational for 30 years.')
```

Describe the 3 code changes you made below in plain English:

<!-- grade_id: cell-f46f8c8ce34b46ec, points: 1, grade: true, solution: true -->


## Problem 2 - Accessing monthly average temperatures (*9 points*)

The table below presents [monthly average temperatures recorded at the Helsinki Malmi airport](https://www.timeanddate.com/weather/finland/helsinki/climate).

| Month     | Temperature [°C] |
| --------- | :--------------: |
| January   | -3.5             |
| February  | -4.5             |
| March     | -1.0             |
| April     | 4.0              |
| May       | 10.0             |
| June      | 15.0             |
| July      | 18.0             |
| August    | 16.0             |
| September | 11.5             |
| October   | 6.0              |
| November  | 2.0              |
| December  | -1.5             |

In the code cells below, write some Python code that allows users to select a month and have the monthly average temperature printed to the screen.
For example, your code should display the following for the month of March:

```
The average temperature in Helsinki in March is -1.0
```

### Grading for Problem 2

Your score for this problem will be based on

- Having your notebook display the monthly average temperature in a selected month, set by defining the variable `selected_month_index`.
    - **NOTE**: We expect you to use the index value to select the month, not the name of the month (e.g., "May").
- Having it work for all 12 months in the year.
- Describe how your code works in a few sentences of plain English in the Markdown cell above your code
- Including inline comments to clearly explain how the code works between most and/or all lines of your code in its Python cell
- Pushing your script to your GitHub repository for this week's lesson.
- Answering the three questions listed at the end of this problem.

<!-- grade_id: cell-8f55ec0c67488f58, points: 0, grade: true, solution: true -->
### Code description

Describe how your code works here.

<!-- grade_id: cell-4e0b1f7ad0ca4693, points: 0, grade: false, solution: true -->
```python
# We have provided the names of the two lists you should use for this problem.
months = None
average_temp = None
### BEGIN SOLUTION
months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', \
         'October', 'November', 'December']
average_temp = [-3.5, -4.5, -1.0, 4.0, 10.0, 15.0, 18.0, 16.0, 11.5, 6.0, 2.0, -1.5]
selected_month_index = 6
print('The average temperature in Helsinki in', months[selected_month_index], 'is', average_temp[selected_month_index])
### END SOLUTION
```

<!-- grade_id: cell-ff185e790e43145e, points: 2, grade: true, solution: false -->
```python
#Validate the length of two lists are 12
ok_(len(months)==12)
ok_(len(average_temp)==12)
```

<!-- grade_id: cell-fb4b3da63919bbeb, points: 2, grade: true, solution: false -->
```python
#Validate that variable months and average_temp are lists
ok_(isinstance(months, list), 'Variable months is not a list')
ok_(isinstance(average_temp, list), 'Variable average_temp is not a list')
```

<!-- grade_id: cell-ccf0ca5467b49c1b, points: 2, grade: true, solution: false -->
```python
#Validate some list elements
assert_equal(months[3],'April')
assert_equal(months[-1],'December')
assert_equal(average_temp[3],4.0)
assert_equal(average_temp[-1],-1.5)
```

<!-- grade_id: cell-ed263e3d32ee8734, points: 3, grade: true, solution: false -->
```python
#Validate the print statement is correct
#Please set the index to 6, run the code cell and run this test cell.
assert_equal('The average temperature in Helsinki in July is 18.0', \
'The average temperature in Helsinki in '+months[selected_month_index]+' is '+str(average_temp[selected_month_index]))
```

## Problem 3 - More practice using Markdown (*4 points*)

The last task in this week's exercise is to answer some questions in this notebook using Markdown.

**Note**: You may want to read a bit more about [formatting text in Github-flavored Markdown](https://help.github.com/articles/basic-writing-and-formatting-syntax/).

1. Give your responses to the three questions about this week's lesson.
    - Be sure to list the question on the first line of the Markdown cell using Markdown for an unordered list
    - Questions to answer about Lesson 2:
        - What did you like?
        - What did you dislike?
        - What would you change?

2. Just for fun, add an image of an animal that you like along with a short caption giving its name and anything special you might like to add.
You should add an image by linking to a website with the image, or if by uploading an image to your GitHub repository and linking to that (if you are looking for a challenge).
We suggest that you search for images in a repository that includes licencing information such as [Wikimedia Commons](https://commons.wikimedia.org/wiki/Main_Page) or [Pixabay](https://pixabay.com/).
You are, of course, also welcome to upload your own animal images.
You can add it in a new Markdown cell under those used to answer the questions in part 1 of this problem.

### Grading for Problem 3

Your score for this problem will be based on

- Your answers to the three questions that are part of this problem
- Posting an image of a favorite animal using Markdown

<!-- grade_id: cell-ec1cb9fcd924ec5a, points: 1, grade: true, solution: true -->
- What did you like?

This lesson was amazing! Python lists are simply mind blowing!

<!-- grade_id: cell-0c0f89bbc72e46d0, points: 1, grade: true, solution: true -->


<!-- grade_id: cell-25fb5ac29fa3f58d, points: 1, grade: true, solution: true -->


<!-- grade_id: cell-0ef0803466212d71, points: 1, grade: true, solution: true -->
My favorite animal:

![A wolverine](https://cdn.pixabay.com/photo/2016/09/10/23/55/wolverine-1660429_1280.jpg)