# Exercise 2 - Lists and indices

The exercise for this week is meant to help you better understand data types and lists in Python, and practice saving changes to your files using Git and GitHub. You will also get to practice the Markdown syntax. 

Exercise 2 consists of 3 problems that you can find in this notebook. For each problem you need to modify the notebook either by fixing problems in existing code, or by adding your own solutions. Remember to save and commit your changes locally, and push your changes to GitHub after each major change! Regular commits will help you to keep track of your changes (and revert them if needed). Pushing your work to GitHub will ensure that you don't lose any work in case your computer crashes (can happen!).  

## Please note

- **Some of the cells in this Jupyter Notebook are "Read-only" which means that you are not able to edit them.**

- **Some of the code cells contain tests. If these tests fail, you know there is still something wrong with your code.**

- **Some of the code cells raise** `NotImplementedError`. **This error tells you that you need to somehow modify that code cell.**


## Where to find help

- Revise [lesson 2 materials](https://geo-python.github.io/site/lessons/L2)
- Check out [the hints for this week's exercise](https://geo-python.github.io/site/lessons/L2/exercise-2.html) on the course webpage

Students at the University of Helsinki are also encouraged to participate in the exercise sessions and discussion on Slack. 

## Problem 0 - Cloning Exercise repository to JupyterLab (*0 points*)

Before starting to work with the actual problems for this week, you should start a new JupyterLab instance and clone **your own** Exercise 2 repository (e.g. `exercise-2-htenkanen`) into the instance using **Git**. *[See Lesson 2 materials](https://geo-python.github.io/site/lessons/L2/git-basics.html) for instructions on how to use `git clone` to get started.* 

## Problem 1 - Making changes to code and using Git (*3 points*)

Your first task for this week is to make some modifications to the Python code in the cells below to fix the errors it gives once you run the code cells.
The code should allow users to find the operational time of an [FMI observation station](http://en.ilmatieteenlaitos.fi/observation-stations) by setting the `selected_station` variable.
*Don't worry about the case of a station name that is not on the list.*

**The main aim is to print out an informative sentence that tells how long the selected station has been operational, for example** 

```
'The Helsinki Harmaja station has been operational for 30 years.'
```
or
```
'The Helsinki Kaivopuisto station has been operational for 115 years.'
```

However, if you try to run all the code cells under Problem 1, you will run into errors. Your task is to find **3 things to fix** in the code cells under Problem 1. Once you succeed, the final print statement should work, and you should be able to modify the output only by setting a different station name.


### Grading for Problem 1

Your score for this problem will be based on

- Fixing **3 problems in the code** to get the code working as expected
- Committing each change separately using Git, and pushing the changes to GitHub
- Listing the changes you needed to make to get the code working in plain English by modifying the existing text or adding a Markdown cell above each change

Here we create and fill lists of station names and starting years for observation.
**NOTE**: Long lines can be split into multiple lines using the "\\" character.

*These code cells are "Read-only" - you won't be able to modify the input lists, but you still need to run them.*

```python
# Read-only code cell
station_names = ['Harmaja', 'Kaisaniemi', 'Kaivopuisto', 'Kumpula', 'lighthouse', \
                'Malmi airfield', 'Suomenlinna aaltopoiju', 'Vuosaari harbour']
```

```python
# Read-only code cell
station_start_years = [1989, 1844, 1904, 2005, 2003, 1937, 2016, 2012]
```

Below we can define the selected station:

<!-- grade_id: selected_station, points: 0, grade: false, solution: true -->
```python
selected_station = None

### BEGIN SOLUTION
selected_station = 'Harmaja'
### END SOLUTION
```

<!-- grade_id: selected_station_test, points: 1, grade: true, solution: false -->
```python
print(selected_station)
### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal

# Validate that selected_station is an element in the list station_names
ok_(selected_station in station_names, \
    'selected_station = '+str(selected_station)+' is not an element in station_names')
### END HIDDEN TESTS
```

Output:
```
None
<ipython-input-6-71a7d9394a5f> in <module>
      5 # Validate that selected_station is an element in the list station_names
      6 ok_(selected_station in station_names, \
----> 7     'selected_station = '+str(selected_station)+' is not an element in station_names')
      8 ### END HIDDEN TESTS

C:\Hyapp\Anaconda3-2019.3\envs\gis\lib\site-packages\nose\tools\trivial.py in ok_(expr, msg)
     20     """
     21     if not expr:
---> 22         raise AssertionError(msg)
     23 
     24 

AssertionError: selected_station = None is not an element in station_names
```

Next we find the location of the selected station in the list of station names:

```python
# Read-only code cell
station_index = station_names.index(selected_station)
```

Output:
```
ValueError                                Traceback (most recent call last)
<ipython-input-7-85522cac0ddd> in <module>
      1 # Read-only code cell
----> 2 station_index = station_names.index(selected_station)

ValueError: None is not in list
```

Now we can calculate the years the station has been operational:

<!-- grade_id: station_years, points: 0, grade: false, solution: true -->
```python
station_years = 2019 - station_start_years[selected_station]
### BEGIN SOLUTION
station_years = 2019 - station_start_years[station_index]
### END SOLUTION
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-8-c944f44b479a> in <module>
      1 ### BEGIN SOLUTION
----> 2 station_years = 2019 - station_start_years[station_index]
      3 ### END SOLUTION
      4 
      5 station_years = 2019 - station_start_years[selected_station]

NameError: name 'station_index' is not defined
```


Now, we can print the station name and number of years of operation to the screen.
A sample print statement on screen is `'The Helsinki Harmaja station has been operational for 30 years.'` The print statements should work for all listed stations only by modifying the `selected_station` -variable that is defined earlier in this problem.

<!-- grade_id: print_statement, points: 0, grade: false, solution: true -->
```python
print_statement = "The Helsinki " + selected_station + " station has been operational for "+ \
      str(station_years)+ " years.

### BEGIN SOLUTION
print_statement = "The Helsinki " + selected_station + " station has been operational for "+ \
      str(station_years)+ " years."
### END SOLUTION

print(print_statement)
```

Output:
```
File "<ipython-input-9-c3e5fd5d1468>", line 7
    str(station_years)+ " years.
                                ^
SyntaxError: EOL while scanning string literal
```

#### Tests for problem 1

Finally, we can run some tests to see if things are working as they should! These tests may help you in solving the problem. Note, we also run some hidden tests for checking the exercises. If these tests pass, it's not a quarantee of a 100 % correct answer!

<!-- grade_id: station_years_test, points: 1, grade: true, solution: false -->
```python
# importing tools for running some tests
from nose.tools import ok_, assert_equal

#Validate that station_years is a variable of integer type
assert_equal(type(station_years),int)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-10-dc916482fd26> in <module>
      3 
      4 #Validate that station_years is a variable of integer type
----> 5 assert_equal(type(station_years),int)

NameError: name 'station_years' is not defined
```

<!-- grade_id: print_statement_test1, points: 2, grade: true, solution: false -->
```python
#Validate that the print statement is correct for Harmaja station:
# Note! Your print statement should work for all the stations!
assert_equal(print_statement, 'The Helsinki Harmaja station has been operational for 30 years.')

### BEGIN HIDDEN TESTS
selected_station = 'Kaivopuisto'
station_years = 2019 - station_start_years[station_names.index(selected_station)]
assert_equal("The Helsinki Kaivopuisto station has been operational for 115 years.", "The Helsinki " + selected_station + " station has been operational for "+ \
      str(station_years)+ " years.")
### END HIDDEN TESTS
```

Output:
```
C:\Hyapp\Anaconda3-2019.3\envs\gis\lib\unittest\case.py in assertEqual(self, first, second, msg)
    850         """
    851         assertion_func = self._getAssertEqualityFunc(first, second)
--> 852         assertion_func(first, second, msg=msg)
    853 
    854     def assertNotEqual(self, first, second, msg=None):

C:\Hyapp\Anaconda3-2019.3\envs\gis\lib\unittest\case.py in _baseAssertEqual(self, first, second, msg)
    843             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    844             msg = self._formatMessage(msg, standardMsg)
--> 845             raise self.failureException(msg)
    846 
    847     def assertEqual(self, first, second, msg=None):

AssertionError: None != 'The Helsinki Harmaja station has been operational for 30 years.'
```

#### Summarize what you did in problem 1 using markdown

Describe the 3 code changes you made below in plain English (you can, for example, use bullet points here):

<!-- grade_id: final_answer_question_test, points: 0, grade: true, solution: true -->


## Problem 2 - Accessing monthly average temperatures (*5 points*)

The table below presents [monthly average temperatures recorded at the Helsinki Malmi airport](https://www.timeanddate.com/weather/finland/helsinki/climate).

| Month     | Temperature [°C] |
| --------- | :--------------: |
| January   | -3.5             |
| February  | -4.5             |
| March     | -1.0             |
| April     | 4.0              |
| May       | 10.0             |
| June      | 15.0             |
| July      | 18.0             |
| August    | 16.0             |
| September | 11.5             |
| October   | 6.0              |
| November  | 2.0              |
| December  | -1.5             |

In the code cells below, write some Python code that allows users to select a month and have the monthly average temperature printed to the screen.
For example, your code should display the following for the month of March:

```
The average temperature in Helsinki in March is -1.0
```

### Grading for Problem 2

Your score for this problem will be based on

- Having your notebook display the monthly average temperature in a selected month, set by defining the variable `selected_month_index`. **NOTE**: We expect you to use the index value to select the month, not the name of the month (e.g., "May").
- Having it work for all 12 months in the year.
- Describe how your code works in a few sentences of plain English in the Markdown cell above your code
- Including inline comments to clearly explain how the code works between most and/or all lines of your code in its Python cell
- Pushing your script to your GitHub repository.

<!-- grade_id: define_lists, points: 0, grade: false, solution: true -->
```python
# here is the variable which you use to set the selected month:
selected_month_index = None

# Here are also the names of the two lists you should use for this problem. Please don't change the variable names!
months = None
average_temp = None

# Using the lists and their indices, generate the desired print statement:
print_statement = None

### BEGIN SOLUTION
months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', \
         'October', 'November', 'December']
average_temp = [-3.5, -4.5, -1.0, 4.0, 10.0, 15.0, 18.0, 16.0, 11.5, 6.0, 2.0, -1.5]
selected_month_index = 6
print_statement = 'The average temperature in Helsinki in ' + str(months[selected_month_index]) + ' is ' + str(average_temp[selected_month_index])

### END SOLUTION
```

<!-- grade_id: define_lists_test, points: 1, grade: true, solution: false -->
```python
# Check your print statement:
print(print_statement)
### BEGIN HIDDEN TESTS
#Validate some list elements
assert_equal(months[3],'April')
assert_equal(months[-1],'December')
assert_equal(average_temp[3],4.0)
assert_equal(average_temp[-1],-1.5)
### END HIDDEN TESTS
```

Output:
```
The average temperature in Helsinki in July is 18.0
```

#### Tests for problem 2

Let's run some additional tests to see if things are working as they should! These tests may help you in solving the problem. Note, we also run some hidden tests for checking the exercises.

<!-- grade_id: list_length_test, points: 0, grade: true, solution: false -->
```python
#Validate the length of two lists are 12
ok_(len(months)==12)
ok_(len(average_temp)==12)
```

<!-- grade_id: list_type_test, points: 1, grade: true, solution: false -->
```python
#Validate that variable months and average_temp are lists
ok_(isinstance(months, list), 'Variable months is not a list')
ok_(isinstance(average_temp, list), 'Variable average_temp is not a list')
```

<!-- grade_id: print_statement_test, points: 2, grade: true, solution: false -->
```python
#Validate the print statement is correct; 
# Set selected_month_index to correspond with July before running this cell.
# Note! Your code should work with any of the 12 months!
assert_equal(print_statement, 'The average temperature in Helsinki in July is 18.0')

### BEGIN HIDDEN TESTS
selected_station_index = 6
assert_equal('The average temperature in Helsinki in July is 18.0', \
'The average temperature in Helsinki in '+months[selected_month_index]+' is '+str(average_temp[selected_month_index]))
### END HIDDEN TESTS
```

## Problem 3 - More practice using Markdown (*2 points*)

The last task in this week's exercise is to answer some questions and add an image in this notebook using Markdown. **Note**: You may want to read a bit more about [formatting text in Github-flavored Markdown](https://help.github.com/articles/basic-writing-and-formatting-syntax/).


### Grading for Problem 3

Your score for this problem will be based on

- Having answered the three questions below
- Posting an image of a favorite animal using Markdown


#### 1. Give your responses to these three questions about this week's lesson:

    - What did you learn?
    - What was unclear?
    - What would you change?
    
*Please use Markdown lists when answering these questions.*
        

<!-- grade_id: answer_questions, points: 1, grade: true, solution: true -->
- What did you like?

This lesson was amazing! Python lists are simply mind blowing!

#### 2. Adding an image

Just for fun, add an image of an animal that you like along with a short caption giving its name and anything special you might like to add.

You can add an image using an URL, or (if you are looking for a challenge) by uploading an image to your GitHub repository and linking that image file to this notebook .

We suggest that you search for images in a repository that includes licencing information such as [Wikimedia Commons](https://commons.wikimedia.org/wiki/Main_Page) or [Pixabay](https://pixabay.com/). You are, of course, also welcome to upload your own animal images. 

<!-- grade_id: insert_image, points: 1, grade: true, solution: true -->
My favorite animal:

![A wolverine](https://cdn.pixabay.com/photo/2016/09/10/23/55/wolverine-1660429_1280.jpg)