# Exercise 2 - Lists and indices (10 points)

The exercise for this week is meant to help you better understand lists in Python, and practice saving changes to your files using Git and GitHub. You will also get to practice the Markdown syntax. 

Exercise 2 consists of 3 problems. For each problem you need to modify the notebook by adding your own solutions. Remember to save and commit your changes locally, and push your changes to GitHub after each major change! Regular commits will help you to keep track of your changes (and revert them if needed). Pushing your work to GitHub will ensure that you don't lose any work in case your computer crashes (can happen!).

**Please don't change the file name**, i.e. do all your editing in the provided `Exercise-2.ipynb` file. 

- **Exercise 2 is due by the start of the next lesson (9:15 am, 16 September 2020)**.
- **We are [working in pairs](https://geo-python-site.readthedocs.io/en/latest/lessons/L2/why-pairs.html) on this exercise**, and we will only grade the repository of the member of your pair that is responsible for this week's exercise.

## Please note

- The Python code cells where you should make changes will contain text that reads

  ```Python
  # YOUR CODE HERE
  raise NotImplementedError()
  ```
  **You should delete that text and replace it with your code** to answer the related question
- Some of the cells in this Jupyter Notebook are "Read-only" which means that you are not able to edit them.
- Some of the code cells contain tests. If these tests fail (you see an error message), you know there is still something wrong with your code.

## Where to find help

- Review the [materials for Lesson 2](https://geo-python-site.readthedocs.io/en/latest/lessons/L2/overview.html)
- Check out [the hints for this week's exercise](https://geo-python-site.readthedocs.io/en/latest/lessons/L2/exercise-2.html#exercise-2-hints) on the course webpage

Students at the University of Helsinki are also encouraged to participate in the exercise sessions and discussion on Slack. 

## Problem 0 - Cloning Exercise repository to JupyterLab (*0 points*)

Before starting to work with the actual problems for this week, you should start a new JupyterLab instance and clone **your own** Exercise 2 repository (e.g. `exercise-2-htenkanen`) into the instance using **Git**. *Check [the Lesson 2 materials](https://geo-python-site.readthedocs.io/en/latest/lessons/L2/git-basics.html) for instructions on how to get started.* 

## Problem 1 - Making and changing lists, and using Git (*3 points*)

Your first task for this week is to create a list of [FMI observation stations](http://en.ilmatieteenlaitos.fi/observation-stations) and the years in which they began operating.
You will then add some additional stations and operation years to your lists, and finally sort the lists such that the newest station is first in the list of operation years.
You will find the list of stations and the years in which they began operating below.

| FMI station name        | First year of operation |
| ----------------------- | :---------------------: |
| lighthouse              | 2003                    |
| Harmaja                 | 1989                    |
| Suomenlinna aaltopoiju  | 2016                    |
| Kumpula                 | 2005                    |
| Kaisaniemi              | 1844                    |

### Grading for Problem 1

Your score for this problem will be based on

- Creating two lists for the FMI station names and years in which they began operating
- Adding the additional stations and operation years to your lists
- Sorting the lists as directed
- Committing each change separately using Git, and pushing the changes to GitHub
- Listing the changes you needed to make to get the code working in plain English by modifying the existing text or adding a Markdown cell above each change

### Part 1 - Creating the lists (1 point)

In the cell below you should create two lists:

- `station_names` should contain the FMI station names in the order they are given above
- `station_start_years` should contain the first years of operation of the stations, again in the order give above

<!-- grade_id: problem_1_ag_p1_station_lists, points: 0, grade: false, solution: true -->
```python
# Create the lists below
station_names = None
station_start_years = None
### BEGIN SOLUTION
station_names = ['lighthouse', 'Harmaja', 'Suomenlinna aaltopoiju', 'Kumpula', 'Kaisaniemi']
station_start_years = [2003, 1989, 2016, 2005, 1844]
### END SOLUTION
```

Let's now check the lists contain the expected values.

<!-- grade_id: problem_1_ag_p1_station_lists_test1, points: 0.5, grade: true, solution: false -->
```python
# This is a test cell that checks that the list lengths are correct
# Running this cell should not produce any errors
assert len(station_names) == 5, 'The station_names list should have 5 items.'
assert len(station_start_years) == 5, 'The station_start_years list should have 5 items.'
```

<!-- grade_id: problem_1_ag_p1_station_lists_test2, points: 0.5, grade: true, solution: false -->
```python
# This is a test cell that checks that the first item in the lists is correct
# Running this cell should not produce any errors
assert station_names[0] == 'lighthouse', 'The first item in the station_names list should be "lighthouse"'
assert station_start_years[0] == 2003, 'The first item in the station_start_years list should be 2003'
### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal
# Validate some additional list elements
assert_equal(station_names[2], 'Suomenlinna aaltopoiju')
assert_equal(station_names[-1], 'Kaisaniemi')
assert_equal(station_start_years[2], 2016)
assert_equal(station_start_years[-1], 1844)
### END HIDDEN TESTS
```

### Part 2 - Modifying the lists (1 point)

Let's now add a few more stations to our lists.
The stations to add are shown in the table below.

| FMI station name | First year of operation |
| ---------------- | :---------------------: |
| Malmi airfield   | 1937                    |
| Vuosaari harbour | 2012                    |
| Kaivopuisto      | 1904                    |

In the cell below you should add the stations and starting years to your existing `station_names` and `station_start_years` lists in the order they are listed above.

<!-- grade_id: problem_1_ag_p2_modify_lists, points: 0, grade: false, solution: true -->
```python
# Add the additional stations to the lists below

### BEGIN SOLUTION
# Two good options here:
# Option 1 - Adding a second list
station_names += ['Malmi airfield', 'Vuosaari harbour', 'Kaivopuisto']
station_start_years += [1937, 2012, 1904]

# Option 2 - Append
#station_names.append('Malmi airfield')
#station_names.append('Vuosaari harbour')
#station_names.append('Kaivopuisto')
#station_start_years.append(1937)
#station_start_years.append(2012)
#station_start_years.append(1904)
### END SOLUTION
```

And we can now check to see whether the modified lists contain the expected values.

<!-- grade_id: problem_1_ag_p2_modify_lists_test1, points: 0.5, grade: true, solution: false -->
```python
# This is a test cell that checks that the list lengths are correct
# Running this cell should not produce any errors
assert len(station_names) == 8, 'The station_names list should have 8 items.'
assert len(station_start_years) == 8, 'The station_start_years list should have 8 items.'
```

<!-- grade_id: problem_1_ag_p2_modify_lists_test2, points: 0.5, grade: true, solution: false -->
```python
# This is a test cell that checks that the last item in the lists is correct
# Running this cell should not produce any errors
assert station_names[-1] == 'Kaivopuisto', 'The last item in the station_names list should be "Kaivopuisto"'
assert station_start_years[-1] == 1904, 'The last item in the station_start_years list should be 1904'
### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal
# Validate some additional list elements
assert_equal(station_names[5], 'Malmi airfield')
assert_equal(station_names[-2], 'Vuosaari harbour')
assert_equal(station_start_years[5], 1937)
assert_equal(station_start_years[-2], 2012)
### END HIDDEN TESTS
```

### Part 3 - Sorting the lists (1 point)

Finally, we can take our lists and sort them. The goal here is to produce two sorted lists:

- The `station_names` list should be sorted so the stations are in alphabetical order
- The `station_start_years` list should be sorted so that the most recent station starting year is first in the list

<!-- grade_id: problem_1_ag_p3_station_years, points: 0, grade: false, solution: true -->
```python
# Sort the lists as directed below

### BEGIN SOLUTION
station_names.sort()
station_start_years.sort()
station_start_years.reverse()
### END SOLUTION
```

Let's now check that the first value in each sorted list is correct.

<!-- grade_id: problem_1_ag_p3_station_years_test1, points: 1, grade: true, solution: false -->
```python
# This is a test cell that checks that the last item in the lists is correct
# Running this cell should not produce any errors
assert station_names[0] == 'Harmaja', 'The first item in the station_names list should be "Harmaja"'
assert station_start_years[0] == 2016, 'The first item in the station_start_years list should be 2016'
### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal
# Validate some additional list elements
assert_equal(station_names[3], 'Kumpula')
assert_equal(station_names[-1], 'lighthouse')
assert_equal(station_start_years[3], 2003)
assert_equal(station_start_years[-1], 1844)
### END HIDDEN TESTS
```

### Part 4 - Understanding the code (0 points)

Looking at the lists used in Problem 1, what might be a problem with how we have sorted them?
Please answer briefly in the cell below.

<!-- grade_id: problem_1_mg_p4_final_answer_question_test, points: 0, grade: true, solution: true -->


### Part 5 - Python + Google (optional, 0 points)

Python has a function called `zip()` that might be helpful in solving the issue mentioned in Part 4.
If you like, you're free to do some Googling and see whether you can figure out how to use the `zip()` function to ensure that the station data is consistent for both lists when sorting the station starting years in reverse order.
This is a completely optional task.

**NOTE**: If you plan to skip this optional problem, be sure to delete the `raise NotImplementedError()` line in the cell below so that you don't have any errors when we run your notebook!

<!-- grade_id: problem_1_mg_p5_zip_test, points: 0, grade: true, solution: true -->
```python
# Add code below if you'd like to test using the zip() function

```

## Problem 2 - Accessing monthly average temperatures (*5 points*)

The table below presents [monthly average temperatures recorded at the Helsinki Malmi airport](https://www.timeanddate.com/weather/finland/helsinki/climate).

| Month     | Temperature [°C] |
| --------- | :--------------: |
| January   | -3.5             |
| February  | -4.5             |
| March     | -1.0             |
| April     | 4.0              |
| May       | 10.0             |
| June      | 15.0             |
| July      | 18.0             |
| August    | 16.0             |
| September | 11.5             |
| October   | 6.0              |
| November  | 2.0              |
| December  | -1.5             |

In the code cells below, write some Python code that allows users to select a month and have the monthly average temperature printed to the screen.
For example, your code should display the following for the month of March:

```
The average temperature in Helsinki in March is -1.0
```

### Grading for Problem 2

Your score for this problem will be based on

- Having your notebook display the monthly average temperature in a selected month, set by defining the variable `selected_month_index`. **NOTE**: We expect you to use the index value to select the month, not the name of the month (e.g., "May").
- Having it work for all 12 months in the year.
- Describe how your code works in a few sentences of plain English in the Markdown cell above your code
- Including inline comments to clearly explain how the code works between most and/or all lines of your code in its Python cell
- Pushing your script to your GitHub repository.

<!-- grade_id: problem_2_ag_p1_define_lists, points: 0, grade: false, solution: true -->
```python
# here is the variable which you use to set the selected month:
selected_month_index = None

# Here are also the names of the two lists you should use for this problem. Please don't change the variable names!
months = None
average_temp = None

# Using the lists and their indices, generate the desired print statement:
print_statement = None

### BEGIN SOLUTION
months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', \
         'October', 'November', 'December']
average_temp = [-3.5, -4.5, -1.0, 4.0, 10.0, 15.0, 18.0, 16.0, 11.5, 6.0, 2.0, -1.5]
selected_month_index = 6
print_statement = 'The average temperature in Helsinki in ' + str(months[selected_month_index]) + ' is ' + str(average_temp[selected_month_index])

### END SOLUTION
```

<!-- grade_id: problem_2_ag_p1_define_lists_test1, points: 1, grade: true, solution: false -->
```python
# Check your print statement:
print(print_statement)
### BEGIN HIDDEN TESTS
#Validate some list elements
assert_equal(months[3],'April')
assert_equal(months[-1],'December')
assert_equal(average_temp[3],4.0)
assert_equal(average_temp[-1],-1.5)
### END HIDDEN TESTS
```

### Additional tests for problem 2

Let's run some additional tests to see if things are working as they should! These tests may help you in solving the problem. Note, we also run some hidden tests for checking the exercises.

<!-- grade_id: problem_2_ag_p1_define_lists_test2, points: 1, grade: true, solution: false -->
```python
#Validate the length of two lists are 12
assert len(months) == 12, 'Wrong length!'
assert len(average_temp) == 12, 'Wrong length!'
```

<!-- grade_id: problem_2_ag_p1_define_lists_test3, points: 1, grade: true, solution: false -->
```python
#Validate that variable months and average_temp are lists
assert isinstance(months, list), 'Variable months is not a list'
assert isinstance(average_temp, list), 'Variable average_temp is not a list'
```

<!-- grade_id: problem_2_ag_p1_define_lists_test4, points: 2, grade: true, solution: false -->
```python
#Validate the print statement is correct; 
# Set selected_month_index to correspond with July before running this cell.
# Note! Your code should work with any of the 12 months!
assert print_statement == 'The average temperature in Helsinki in July is 18.0'
### BEGIN HIDDEN TESTS
selected_station_index = 6
assert_equal('The average temperature in Helsinki in July is 18.0', \
'The average temperature in Helsinki in '+months[selected_month_index]+' is '+str(average_temp[selected_month_index]))
### END HIDDEN TESTS
```

## Problem 3 - More practice using Markdown (*2 points*)

The last task in this week's exercise is to answer some questions and add an image in this notebook using Markdown. **Note**: You may want to read a bit more about [formatting text in Github-flavored Markdown](https://help.github.com/articles/basic-writing-and-formatting-syntax/).


### Grading for Problem 3

Your score for this problem will be based on

- Having answered the three questions below
- Posting an image of a favorite animal using Markdown


#### 1. Give your responses to these three questions about this week's lesson:

    - What did you learn?
    - What was unclear?
    - What would you change?
    
*Please use Markdown lists when answering these questions.*
        

<!-- grade_id: problem_3_mg_p1_questions, points: 1, grade: true, solution: true -->
- What did you like?

This lesson was amazing! Python lists are simply mind blowing!

#### 2. Adding an image

Just for fun, add an image of an animal that you like along with a short caption giving its name and anything special you might like to add.

You can add an image using an URL, or (if you are looking for a challenge) by uploading an image to your GitHub repository and linking that image file to this notebook .

We suggest that you search for images in a repository that includes licencing information such as [Wikimedia Commons](https://commons.wikimedia.org/wiki/Main_Page) or [Pixabay](https://pixabay.com/). You are, of course, also welcome to upload your own animal images. 

<!-- grade_id: problem_3_mg_p1_insert_image, points: 1, grade: true, solution: true -->
My favorite animal:

![A wolverine](https://cdn.pixabay.com/photo/2016/09/10/23/55/wolverine-1660429_1280.jpg)